
# Electrónica II. Filtros activos

Este notebook es un ejemplo de una prueba de evaluación continua. Se tratan algunos ejercicios básicos de electrónica, que el alumno deberá completar para ser puntuado posteriormente.

![Python_logo](https://www.python.org/static/community_logos/python-logo-master-v3-TM.png)

## Ejercicio 1. Cálculo del orden de un Filtro de Butterworth (3 puntos)

Un filtro es un circuito que permite el paso de señales de unas determinadas frecuencias y atenúa (o impide) el resto. Según el factor de calidad, existen una serie de filtros definidos. Uno de ellos es el filtro de Butterworth.

Escriba una función que, recibiendo los parámetros del filtro: $A_{min},   A_{max}   ,  k$ devuelva el orden del filtro de Butterworth.

Asegúrese de que tiene en cuenta los casos en los que $A_{min}<=A_{max}$, arrojando un error con `ValueError`.


```python
import numpy as np
from math import ceil
```

<!-- grade_id: ej1_answer, points: 0, grade: false, solution: true -->
```python
def butter(Amin,Amax,k):
    "Con la atenuación mínima de la banda suprimida (Amin), la atenuación "
    "máxima de la banda pasante (Amax), y la selectividad del filtro (k),"
    "calcule el orden del filtro de Butterworth correspondiente."
    
    ### BEGIN SOLUTION
    if Amin <= Amax:
        raise ValueError("Amin debe ser mayor que Amax")
    D=(np.power(10, 0.1*Amin)-1)/(np.power(10, 0.1*Amax)-1)
    n=ceil(np.log(D)/(2*np.log(1/k)))
    return n
    ### END SOLUTION
```

Su función debería imprimir `1` para los valores `[12, 3, 0.053]`, respectivamente. Compruébelo ejecutando la siguiente celda:

```python
butter(12,3,0.053)
```

<!-- grade_id: ej1_testing, points: 1.5, grade: true, solution: false -->
```python
"Se comprueba que la función retorna el valor correcto para distintas entradas"

assert butter(45,3,0.2) == 4
assert butter(10,3,0.8) == 5
assert butter(27,3,0.36) == 4

### BEGIN HIDDEN TESTS
assert butter(50,3,0.099) == 3
assert butter(14,3,0.759) == 6
### END HIDDEN TESTS
```

<!-- grade_id: ej1_incorrect_testing, points: 1.5, grade: true, solution: false -->
```python
"Se comprueba que la función lanza un error para entradas incorrectas"

try:
    butter(3,12,0.053)
except ValueError:
    pass
else:
    raise AssertionError("¡No saltó ningún error!")

### BEGIN HIDDEN TESTS
try:
    butter(3,3,0.053)
except ValueError:
    pass
else:
    raise AssertionError("¡No saltó ningún error!")
### END HIDDEN TESTS
```

## Ejercicio 2. Filtro Paso Bajo **$1^{er}$**orden (1.5 puntos)

Escriba en formato LaTeX correctamente la fórmula del filtro paso bajo de primer orden


<!-- grade_id: ej2_answer, points: 1.5, grade: true, solution: true -->
\begin{equation}
F(s) = A_0 * \frac{1}{1+\tau s} = A_0 * \frac{1}{1+\frac{s}{w_c}}
\end{equation}


## Ejercicio 3. Filtro activo vs. pasivo (1.5 punto)

Indique las diferencias entre un filtro activo y uno pasivo.

**Bonus:** ¿Cómo debe ser un filtro para ser considerado ideal? **(+0.5 pts)**

<!-- grade_id: ej3_answer, points: 1.5, grade: true, solution: true -->
Hay dos tipos de filtros:

- Pasivos: únicamente atenúan la señal para algunas frecuencias. Se construyen con R, L, C. Pueden ser usados en frecuencias altas.
- Activos: utilizan componentes que pueden presentar ganancia como el Amplificador Operacional (AO). Se suelen usar hasta frecuencias del orden de los MHz, debido a las limitaciones de sus componentes.

Un filtro ideal no debe atenuar ni desfasar la señal en la banda pasante (gama de frecuencias que no son alteradas), y debe tener ateuación infinita en la banda rechazada (frecuencias filtradas).

## Ejercicio 4. Respuesta en frecuencia (4 puntos)

Represente gráficamente la respuesta en frecuencia de un filtro de Butterworth Paso Bajo, de frecuencia de corte 1 rad/s. Compare los filtros si son de orden 1,2,3...9,10. Como ayuda, puede utilizar las siguientes funciones:
- scipy.signal
- matplotlib.pyplot

Explique qué características tiene un filtro de Butterworth.