
# Matemáticas I. Conceptos básicos

En este primer notebook, vamos a tratar algunos ejemplos básicos como toma de contacto con el entorno.

![](jupyter.png) 
![Python_logo](https://www.python.org/static/community_logos/python-logo-master-v3-TM.png)

## Ejercicio 1. Sumatorio (3 puntos)

Escriba una función que reciba un parámetro $n$ y retorne el valor $res$, de tal forma que:

\begin{equation*}
res = \sum_{k=1}^n \left ( k^2  \right)
\end{equation*}

Asegúrese de que tiene en cuenta los casos en los que $n<1$, arrojando un error con `ValueError`.

<!-- grade_id: ej1_answer, points: 0, grade: false, solution: true -->
```python
def sumatorio(n):
    " Realice el sumatorio de los cuadrados desde 1 hasta n "
    
    ### BEGIN SOLUTION
    if n < 1:
        raise ValueError("n debe ser mayor o igual a 1")
    res=0
    for val in range (n+1):
        res+=val*val
    return res
    ### END SOLUTION
```

Su función debería imprimir `285` para $n$ =`9`. Compruébelo ejecutando la siguiente celda:

```python
sumatorio(9)
```

<!-- grade_id: ej1_testing, points: 1.5, grade: true, solution: false -->
```python
"Se comprueba que la función retorna el valor correcto para distintas entradas"

assert sumatorio(1) == 1
assert sumatorio(4) == 30
assert sumatorio(5) == 55
assert sumatorio(10) == 385

### BEGIN HIDDEN TESTS
assert sumatorio(25) == 5525
assert sumatorio(37) == 17575
### END HIDDEN TESTS
```

<!-- grade_id: ej1_incorrect_testing, points: 1.5, grade: true, solution: false -->
```python
"Se comprueba que la función lanza un error para entradas incorrectas"

try:
    sumatorio(0)
except ValueError:
    pass
else:
    raise AssertionError("¡No saltó ningún error!")

### BEGIN HIDDEN TESTS
try:
    sumatorio(-8)
except ValueError:
    pass
else:
    raise AssertionError("¡No saltó ningún error!")
### END HIDDEN TESTS
```

## Ejercicio 2. Exponente (3 puntos)

Escriba una función que cumpla los siguientes requisitos:

- Recibe un parámetro $num$
- Devuelve el exponente de la potencia de 2 que, siendo el resultado menor o igual a $num$, es el más grande.

Por ejemplo: `exp(35)` devuelve $5$ ya que $2^5$ es menor que $35$, y $2^6$ se pasa

Tenga en cuenta los casos en los que la función no se puede aplicar

<!-- grade_id: ej2_answer, points: 0, grade: false, solution: true -->
```python
def exp(num):
    " Calcule el exponente de la mayor potencia de 2, siendo el "
    " resultado menor que num "
    
    ### BEGIN SOLUTION
    if num < 1:
        raise ValueError("n debe ser mayor o igual a 1")
    val=0
    result=0
    while num>=val:
        val=2**(result+1)
        result=result+1
    return result-1
    ### END SOLUTION
```

Su función debería imprimir $5$ para un valor de entrada de $35$. Compruébelo ejecutando la siguiente celda:

```python
exp(35)
```

<!-- grade_id: ej2_testing, points: 1.5, grade: true, solution: false -->
```python
"Se comprueba que la función retorna el valor correcto para distintas entradas"

assert exp(1) == 0
assert exp(12) == 3
assert exp(57) == 5
assert exp(64) == 6

### BEGIN HIDDEN TESTS
assert exp(194) == 7
assert exp(1024) == 10
assert exp(31456) == 14
### END HIDDEN TESTS
```

<!-- grade_id: ej2_incorrect_testing, points: 1.5, grade: true, solution: false -->
```python
"Se comprueba que la función lanza un error para entradas incorrectas"

try:
    exp(0)
except ValueError:
    pass
else:
    raise AssertionError("¡No saltó ningún error!")

### BEGIN HIDDEN TESTS
try:
    exp(-54)
except ValueError:
    pass
else:
    raise AssertionError("¡No saltó ningún error!")
### END HIDDEN TESTS
```

## Ejercicio 3. **$2^a$** Ley de Newton (1 puntos)

Escriba en formato LaTeX correctamente la fórmula de la segunda Ley de Newton. No olvide colocar vector donde corresponda.

**Bonus:** añada una breve explicación para obtener puntuación extra **(+0,5 pts)**


<!-- grade_id: ej3_answer, points: 1, grade: true, solution: true -->
$$ \sum \vec{F} = m\vec{a} $$

La segunda ley de Newton establece que si una fuerza neta es aplicada en un objeto, la velocidad del objeto cambiará dado que su dirección o rapidez cambiará.

## Ejercicio 4. Representación gráfica (3 puntos)

Utilice la librería `matplotlib` para representar gráficamente la recta que pasa por los 5 puntos que usted elija.

Explique por qué ha escogido esos puntos. ¿Qué características tienen que tener los puntos para conformar una recta? Escriba la ecuación de la recta: `y=ax+b`

**Bonus:** indique el nombre de los ejes cartesianos para obtener puntuación extra **(+0,5 pts)**

```python

```